# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
from pathlib import Path
HOME = str(Path.home())
ALLERGIES_FILE="/data/allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
### BEGIN SOLUTION
def allergy_count(json_file):
    with open(json_file) as f:
        allergies = json.load(f)
        
    return(len(allergies.get('entry')))
### END SOLUTION
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
### BEGIN SOLUTION
def patient_count(json_file):
    patients = set()
    with open(json_file) as f:
        allergies = json.load(f)
        
    for entry in allergies.get('entry'):
        resource = entry.get('resource')
        patient = resource.get('patient')
        name = patient.get('display')
        patients.add(name)
        
    return patients
### END SOLUTION
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut', 'Paul Boal'}
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
### BEGIN SOLUTION
def allergy_per_patient(json_file):
    patients = {}
    with open(json_file) as f:
        allergies = json.load(f)
        
    for entry in allergies.get('entry'):
        resource = entry.get('resource')
        patient = resource.get('patient')
        name = patient.get('display')
        patients[name] = patients.setdefault(name,0) + 1
        
    return patients
### END SOLUTION
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
def allergy_list(json_file): 
    """(str)-> list
    This function takes a file and reads the JSON from the file and returns the name,
    allergy and reaction for each entry."""
    
    patients = []
    with open(json_file) as f:
        allergies = json.load(f)
        
    for item in allergies.get('entry'):
        name = item.get('resource').get('patient').get('display')
        reaction = item.get('resource').get('reaction')[0]['manifestation'][0]['text']
        allergy = item.get('resource').get('substance').get('text')
        patients.append([name,allergy,reaction])
        
        
    return patients
### END SOLUTION
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list(ALLERGIES_FILE) == output

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
def allergy_reaction(json_file,patient,substance):
    """(str,str,str)->str
    This function takes a file, a patient name and a substance then reads the JSON from the file.
    It then takes the patients name and the substance, and returns the reaction that the patient
    would experience if they took the substance."""
      
    for item in allergy_list(json_file):
        if item[0:2]==[patient,substance]:
            return item[-1]
### END SOLUTION
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---
---

# Stretch (Extra) Problems

Work on either of the stretch problems below can earn you up to 25 free points toward the midterm assignment.  That is, if you complete one of these extra problems successfully, you can skip 1 of the problems that will appear on the midterm exam coming up next week.

The midterm will be distribute this Saturday 3/13.

This assignment is due on Sunday 3/14.  If you are trying for one of these extra problems Slack me, and I'll provide you feedback on how you did on these before end of day Monday 3/15.  That way you can choose what to complete on the midterm.


---
---

### STRETCH for March 2021 - For those looking for an additional challenge

As I've mentioned in class, CMS is now enforcing a rule around price transparency.  Every facility that take Medicare payments is required to publish a "machine readable" file with it's pricing infomration for a number of common procedures across all of the payers they work with.  There are two examples of such files in the `/data/` directory: `whiteriver.json` and `saline.xml`.

If you want to compare contracted prices across these two hospitals, you'll need to read in the information from both of those files into some kind of data structure, then merge the data together from those two files.  See what you can do.

See if you can create an output file that has the following fields:
* HOSPITAL
* PROCEDURE_CODE
* PAYER
* AMOUNT

If you choose to work on this, you may get stuck at some point and you won't know if you're _doing it right_. Make some assumptions. Document your questions in this notebook.



```python
import pandas as pd
from pandas.io.json import json_normalize
import json
import xml.etree.ElementTree as ET

whiteriver_file ="/data/whiteriver.json"
saline_file = "/data/saline.xml"
```

```python
#Whiteriver Data
#Takes json file and converts to dataframe
with open(whiteriver_file) as f:
        whiteriver = json.load(f)
wr_df = json_normalize(whiteriver['root']['StandardCharges'])

#Chooses list of columns to keep
cols = list(wr_df.columns)
payer = cols[13:] #I had to physically look at this list. Is there a way to do this blind?
wr_df = pd.melt(wr_df,id_vars = ['ProcedureCode'],value_vars = payer, var_name = 'PAYER',value_name = 'AMOUNT')

#Renames,adds and reorders columns
wr_df = wr_df.rename(columns = {'ProcedureCode':'PROCEDURE_CODE'})
wr_df['HOSPITAL'] = 'White River'
wr_df = wr_df[['HOSPITAL','PROCEDURE_CODE', 'PAYER', 'AMOUNT']]
wr_df.head()
```

Output:
```
HOSPITAL PROCEDURE_CODE                         PAYER  AMOUNT
0  White River          00170  AARPMedicareManagedCare_OPPS       0
1  White River          00220  AARPMedicareManagedCare_OPPS       0
2  White River          01996  AARPMedicareManagedCare_OPPS       0
3  White River          10004  AARPMedicareManagedCare_OPPS       0
4  White River          10005  AARPMedicareManagedCare_OPPS  609.74
```

```python
#Saline Data
#Creates a list of lists containing each item code, payers,
# and charge for each payer
tree = ET.parse(saline_file)
root = tree.getroot()
item = []

for child in root.iter('Item'):
    for contract in child.iter('Contract'):
        code = child.attrib['Code']
        payer = contract.attrib['Payer']
        amount = contract.attrib['Charge']
        item.append([code,payer,amount])
```

```python
#Creates a dataframe with Saline data, adds hospital name and reorders columns
saline_df = pd.DataFrame(item,columns = ['PROCEDURE_CODE','PAYER','AMOUNT'])
saline_df['HOSPITAL'] = 'Saline Memorial'
saline_df = saline_df[['HOSPITAL','PROCEDURE_CODE','PAYER','AMOUNT']]
saline_df.head()
```

Output:
```
HOSPITAL PROCEDURE_CODE  \
0  Saline Memorial            139   
1  Saline Memorial            153   
2  Saline Memorial            175   
3  Saline Memorial            175   
4  Saline Memorial            175   

                                            PAYER        AMOUNT  
0                          HUMANA INC. - (HMO-MR)   7997.700000  
1  UNITED HEALTHCARE INSURANCE COMPANY - (HMO-MR)   4372.770000  
2                          HUMANA INC. - (HMO-MR)   8850.580000  
3  UNITED HEALTHCARE INSURANCE COMPANY - (HMO-MR)   6377.260000  
4     UNITED HEALTHCARE INSURANCE COMPANY - (POS)  11306.760000
```

```python
#Combines whiteriver and saline dataframes,and changes amount column to int type
combined_df = pd.concat([wr_df,saline_df],ignore_index = True)
combined_df['AMOUNT'] = combined_df['AMOUNT'].apply(lambda x: int(float(x)))
combined_df
```

Output:
```
HOSPITAL PROCEDURE_CODE  \
0          White River          00170   
1          White River          00220   
2          White River          01996   
3          White River          10004   
4          White River          10005   
...                ...            ...   
10611  Saline Memorial          U0002   
10612  Saline Memorial          U0002   
10613  Saline Memorial          U0002   
10614  Saline Memorial          U0002   
10615  Saline Memorial          U0002   

                                                 PAYER  AMOUNT  
0                         AARPMedicareManagedCare_OPPS       0  
1                         AARPMedicareManagedCare_OPPS       0  
2                         AARPMedicareManagedCare_OPPS       0  
3                         AARPMedicareManagedCare_OPPS       0  
4                         AARPMedicareManagedCare_OPPS     609  
...                                                ...     ...  
10611                   ARKANSAS MEDICAID - (Medicaid)      51  
1
```

```python
#Since I've been thinking about doing this topic for my final project, I've thought a little about this data. I've looked at about 20 hospitals 
#in Missouri and have seen a great deal of variety of formats and content, and very few of those actually had any payer information. Since each 
#one I've looked at is so different, including these 2 examples, I've been unsure what I can automate and create functions for since each one is
#so unique. Are there good ways to create functions for some of the tasks above or is it best to manually clean each file?
```

---
---

### STRETCH from March 2020 - For those looking for an additional challenge

The Coronavirus is creating quite the stir right now.  There are some sources suggesting that trends show it is going to be significantly more serious than SARS was back in the 2002 timeframe.  Here's one visualization trying to demonstrate that: https://www.reddit.com/r/China_Flu/comments/ev2b4v/i_updated_some_charts_comparing_this_outbreak/

Someone on Kaggle has generously already compiled a dataset based on information from Johns Hopkins about the Coronavirus outbreak.  https://www.kaggle.com/brendaso/2019-coronavirus-dataset-01212020-01262020  Create a Kaggle account, if you don't already have one.  Download this data set and then upload it to your Jupyter Home folder.  (The "up arrow" button is for uploading a file.)

Use Python's built-in `csv` module to read the data from this file and generate the following information: **what are the total confirmed cases in all of Mainland China as of the latest information in the data set?**  Some important things to note:
* Each entry for a given city has the **cumulative** number of cases.  So that column is not additive (it cannot be summed).  You'll have to find a way to filter your data for the last day for each city, then total those up.
* If you choose to parse the date column, you will want to lookup how to do that using Python's `datetime` module.  Especially the `strptime` function.  https://docs.python.org/3/library/datetime.html#strftime-strptime-behavior  Hint: you can parse a date string in the format 2/17/2020 using the code below.  This link will tell you what things like `%m` and `%Y` mean.

```
from datetime import datetime
d = datetime.strptime('2/17/2020', '%m/%d/%Y')
```

If you want to take this another step, **create a list of tuples that contain (observate date, total confirmed) totalled over all locations represented in the data**

```python

```

```python

```

---

## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Save this note with Ctrl-S (or Cmd-S)
2. Skip down to the last command cell (the one starting with `%%bash`) and run that cell.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

```python
assert False, "DO NOT REMOVE THIS LINE"
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-22-cef8011cb395> in <module>
----> 1 assert False, "DO NOT REMOVE THIS LINE"

AssertionError: DO NOT REMOVE THIS LINE
```

---

```python
%%bash
git pull
git add week06_assignment_2.ipynb
git commit -a -m "Submitting the week 6 programming assignment"
git push
```

Output:
```
Already up to date.
On branch main
Your branch is up to date with 'origin/main'.

Untracked files:
	../week02/week02_inclass.ipynb
	../week04/week04_lookups.ipynb

nothing added to commit but untracked files present
Everything up-to-date
```


---

If the message above says something like _Submitting the week 6 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.

```python

```